Absence of electron transfer-associated changes in the time-dependent X-ray free-electron laser structures of the photosynthetic reaction center

Author Response

The following is the authors’ response to the previous reviews.

Reviewer #1 Public Review

“First, I agree with the authors of this manuscript that conformational changes in the XFEL structures with 2.8 A resolution are not reliable enough for demonstrating the subtle changes in the electron transfer events in this bacterial photosynthesis system. Actually, the data statistics in the paper by Dods et al. showed that the high-resolution range of some of the XFEL datasets may include pretty high noise (low CC1/2 and high Rsplit) so the comparison of the subtle conformational changes of the structures is problematic.

The manuscript by Gai Nishikawa investigated time-dependent changes in the energetics of the electron transfer pathway based on the structures by Dods et al. by calculating redox potential of the active and inactive branches in the structures and found no clear link between the time-dependent structural changes and the electron transfer events in the XFEL structures published by Dods, R.et al. (2021). This study provided validation for the interpretation of the structures of those electrontransferring proteins.

The paper was well prepared.”

Thank you very much for your positive and insightful comment. We greatly appreciate your suggestion regarding the high noise levels of the XFEL structures. Including this information in the Introduction section will draw readers’ attention to the concerns about the reliability of these XFEL structures. We have incorporated it into the Introduction section.

Reviewer #2 Public Review

“The manuscript by Nishikawa et al. addresses time-dependent changes in the electron transfer energetics in the photosynthetic reaction center from Blastochloris viridis, whose time-dependent structural changes upon light illumination were recently demonstrated by time-resolved serial femtosecond crystallography (SFX) using X-ray free-electron laser (XFEL) (Dods et al., Nature, 2021). Based on the redox potential Em values of bacteriopheophytin in the electron transfer active branch (BL) by solving the linear Poisson-Boltzmann equation, the authors found that Em(HL) values in the charge-separated 5-ps structure obtained by XFEL are not clearly changed, suggesting that the P+HL- state is not stabilized owing to protein reorganization. Furthermore, chlorin ring deformation upon HL- formation, which was expected from their QM/MM calculation, is not recognized in the 5ps XFEL structure. Then the authors concluded that the structural changes in the XFEL structures are not related to the actual time course of charge separation. They argued that their calculated changes in Em and chlorin ring deformations using the XEFL structures may reflect the experimental errors rather than the real structural changes; they mentioned this problem is due to the fact that the XFEL structures were obtained at not high resolutions (mostly at 2.8 Å). I consider that their systematic calculations may suggest a useful theoretical interpretation of the XFEL study. However, the present manuscript insists as a whole negatively that the experimental errors may hamper to provide the actual structural changes relevant to the electron transfer events.”

Thank you for your feedback on our manuscript. We appreciate your positive assessment of our systematic calculations and theoretical interpretation of the XFEL study. We have carefully considered your comments and made the necessary revisions to address your concerns.

Reviewer #2 Recommendations for the authors

“The authors have satisfied my concerns mostly, in particular by providing the Em(QA) changes, which seem to be more attractive in the present form. However, the Em(QA) value(s), at least in the dark structure, should be provided, and the procedure of the calculation for the Em(QA) value(s) should be described in METHODS "Calculation of Em".

The calculated Em(QA) values for dataset a and dataset b in the dark structure are –223 mV and – 209 mV, respectively, using the reference Em value of –256 mV versus NHE for menaquinone-2 in water [Photosynth. Res. 134 (2017) 193]. These calculated values are comparable to experimentally measured values of –150 mV for PbRC from Blastochloris viridis (naphtoquinone) [Biochim. Biophys. Acta 440 (1976) 622] and –180 mV for PbRC from Rhodobacter sphaeroides (ubiquinone) [Arch. Biochem. Biophys 172 (1976) 329].

We have now provided this information in the Method (“Calculation of Em”) and Results and Discussion (“Relevance of structural changes observed in XFEL structures”) sections.

Changes from the earlier version to the revised version

Abstract

Abstract

Using the X-ray free-electron laser (XFEL) structures of the photosynthetic reaction center from Blastochloris viridis that show light-induced time-dependent structural changes [Dods, R.et al. (2021) Nature 589, 310-314], we investigated time-dependent changes in the energetics of the electron transfer pathway, considering the entire protein environment of the protein structures and titrating the redox active sites in the presence of all fully equilibrated titratable residues. In the dark and charge-separation intermediate structures, the calculated redox potential (Em) values for the accessory bacteriochlorophyll and bacteriopheophytin in the electron-transfer active branch (BL and HL) are higher than those in the electron-transfer inactive branch (BM and HM). However, the stabilization of the [PLPM]•+HL•– state owing to protein reorganization is not clearly observed in the Em(HL) values in the charge-separated 5-ps ([PLPM]•+HL•– state) structure. Furthermore, the expected chlorin ring deformation upon formation of HL•–(saddling mode) is absent in the HL geometry of the original 5-ps structure. These findings suggest that there is no clear link between the time-dependent structural changes and the electron transfer events in the XFEL structures.

Introduction

Photosynthetic reaction centers from purple bacteria (PbRC) are heterodimeric reaction centers, which are formed by the protein subunits L and M (Figure 1). In PbRC from Blastochloris viridis, the electronic excitation of the bacteriochlorophyll b (BChlb) pair, [PLPM], leads to electron transfer to accessory BChlb, BL, followed by electron transfer via bacteriopheophytin b (BPheob), HL, to menaquinone, QA, along the electron-transfer active L branch (A branch) 1. Electron transfer further proceeds from QA to ubiquinone, QB, which is coupled with proton transfer via charged and polar residues in the QB binding region 2. Although the counterpart M branch (B branch) is essentially electron-transfer inactive, mutations of the Phe-L181/Tyr-M208 pair to tyrosine/phenylalanine lead to an increase in the yield of [PLPM]•+HM•- formation (∼30%), which suggests that these residues are responsible for the energetic asymmetry in the electron transfer branches (e.g., 3). The anionic states BL •-, HL•-, and QA•- form in ∼3.5 ps, ∼5 ps, and ∼200 ps upon the formation of the electronically excited [PLPM]* state, respectively 4. The anionic state formation induces not only reoriganization of the protein environment 5 but also out-of-plane distortion of the chlorin ring 6. Indeed, two distinct conformations of HL•- were reported in spectroscopic studies of PbRC from Rhodobacter sphaeroides 7.

Electron transfer pathways along the L- and M-branches in PbRC from Blastochloris viridis. The PbRC is composed of the L (red), M (blue), H (gold), and C (yellow) subunits. [PLPM]: BChlb pair; BL and BM: accessory BChlb; HL and HM: BPheob; QA: primary quinone (menaquinone); Fe: non-heme Fe complex.

Earlier version: Recently, using the X-ray free electron laser (XFEL), light-induced electron density changes and structural changes of PbRC were analyzed at 1 ps, 5 ps, 20 ps, 300 ps, and 8 μs upon the electronic excitation of [PLPM] at 960 nm 8: the 1 ps XFEL structure represents the [PLPM]* state, the 5 ps and 20 ps XFEL structures represent the charge-separated [PLPM]•+H •– state, and the 300 ps and 8 μs XFEL structures represent the charge-separated [PLPM]•+QA •– state. According to Dods et al. 8, these XFEL structures revealed how the charge-separation process was stabilized by protein conformational dynamics. However, the conclusions drawn from these XFEL structures are based on data with limited resolution. Specifically, 8 out of 9 XFEL structures have a resolution of 2.8 Å (atomic coordinates from PDB codes: 5O4C, 6ZI4, and 6ZI5 for dataset a and 6ZHW, 6ZID, 6ZI6, 6ZI9, and 6ZIA for dataset b) 8. Furthermore, the data statistics provided by Dods et al. indicate that the high-resolution range of some XFEL datasets exhibits high levels of noise, as evidenced by low CC1/2 and high Rsplit values. These observations raise concerns about the reliable comparison of subtle conformational changes among these XFEL structures. Hence, caution must be exercised when interpreting these XFEL structures in terms of their ability to accurately capture relevant conformational changes.

Revised version: Recently, using the X-ray free electron laser (XFEL), light-induced electron density changes and structural changes of PbRC were analyzed at 1 ps, 5 ps, 20 ps, 300 ps, and 8 μs upon the electronic excitation of [PLPM] at 960 nm 8: the 1 ps XFEL structure represents the [PLPM]* state, the 5 ps and 20 ps XFEL structures represent the charge-separated [PLPM]•+HL•– state, and the 300 ps and 8 μs XFEL structures represent the charge-separated [PLPM]•+QA•– state. According to Dods et al. 8, these XFEL structures revealed how the charge-separation process was stabilized by protein conformational dynamics. However, the conclusions drawn from these XFEL structures are based on data with limited resolution. Specifically, 8 out of 9 XFEL structures have a resolution of 2.8 Å (atomic coordinates from PDB codes: 5O4C, 6ZI4, and 6ZI5 for dataset a and 6ZHW, 6ZID, 6ZI6, 6ZI9, and 6ZIA for dataset b) 8. The data statistics may indicate that the high-resolution range of some XFEL datasets exhibits high levels of noise (e.g., low CC1/2). These observations raise concerns about the reliable comparison of subtle conformational changes among these XFEL structures. Hence, caution must be exercised when interpreting these XFEL structures in terms of their ability to accurately capture relevant conformational changes.

Here, we investigated how the redox potential (Em) values of the BChlb and BPheob cofactors for one-electron reduction change as electron transfer proceeds using the dark (0 ps), 1 ps, 5 ps, 20 ps, 300 ps, and 8 μs XFEL structures, solving the linear Poisson-Boltzmann equation, and considering the protonation states of all titratable sites in the entire protein. Structural changes (e.g., side-chain orientation) in the protein environment can be analyzed in the Em shift, as Em is predominantly determined by the sum of the electrostatic interactions between the redox-active site and all other groups (i.e., residues and cofactors) in the protein structure. Subtle structural changes of the BChlb and BPheob chlorin rings, which may not be pronounced even in the Em shift 6, can be analyzed in the out-of-plane distortion of the chlorin rings using a normal-coordinate structural decomposition (NSD) analysis 9, 10 with a combination of a quantum mechanical/molecular mechanical (QM/MM) approach in the entire PbRC protein environment.

Methods

Coordinates and atomic partial charges

The atomic coordinates of PbRC from Blastochloris viridis were taken from the XFEL structures determined at 0 ps (dark state; PDB code 5O4C for dataset a and 5NJ4 for dataset b), 1 ps ([PLPM]* state; PDB code, 6ZHW for dataset b), 5 ps ([PLPM] •+HL•– state; PDB code, 6ZI4 for dataset a and 6ZID for dataset b), 20 ps ([PLPM] •+HL•– state; PDB code, 6ZI6 for dataset b), 300 ps ([PLPM] •+QA•– state; PDB code, 6ZI5 for dataset a and 6ZI9 for dataset b), and 8 μs ([PLPM] •+QA•- state; PDB code, 6ZIA for dataset b). Atoms with 30% occupancy for the photoactivated state 8 were used wherever present. Hydrogen atoms were generated and energetically optimized with CHARMM 11. The atomic partial charges of the amino acids were obtained from the all-atom CHARMM22 12 parameter set. For diacylglycerol, the Fe complex 13, and menaquinone 14, the atomic charges were adopted from previous studies. The atomic charges of BChlb and BPheob (BChlb, BChlb•+, BChlb•–, BPheob, and BPheob•–) were determined by fitting the electrostatic potential in the neighborhood of these molecules using the RESP procedure 15 (Tables S1). The electronic densities were calculated after geometry optimization using the DFT method with the B3LYP functional and 6-31G** basis sets in the JAGUAR program 16. For the atomic charges of the nonpolar CHn groups in the cofactors (e.g., the phytol chains of BChlb and BPheob and the isoprene side chains of quinone), a value of +0.09 was assigned to nonpolar H atoms.

Calculation of Em: solving the linear Poisson-Boltzmann equation

Earlier version: The Em values in the protein were determined by calculating the electrostatic energy difference between the two redox states in a reference model system. This was achieved by solving the linear Poisson-Boltzmann equation with the MEAD program 17 and using Em(BChlb) = –665 mV and Em(BPheob) = –429 mV (based on Em(BChlb) = –700 mV and Em(BPheob) = –500 mV for one-electron reduction measured in dimethylformamide 18, 19), considering the solvation energy difference). The difference in the Em value of the protein relative to the reference system was added to the known Em value. To account for the ensemble of protonation patterns, a Monte Carlo method with Karlsberg was used for sampling 20. The linear Poisson-Boltzmann equation was solved using a three-step grid-focusing procedure with resolutions of 2.5 Å, 1.0 Å, and 0.3 Å. Monte Carlo sampling provided the probabilities [Aox] and [Ared] of the two redox states of molecule A, and Em was evaluated using the Nernst equation. A bias potential was applied to ensure an equal amount of both redox states ([Aox] = [Ared]), thus determining the redox midpoint potential as the resulting bias potential. To ensure consistency with previous computational results, we used identical computational conditions and parameters as previous studies (e.g., 13), performing all computations at 300 K, pH 7.0, and an ionic strength of 100 mM. The dielectric constants were set to 4 for the protein interior and 80 for water.

Revised version: The Em values in the protein were determined by calculating the electrostatic energy difference between the two redox states in a reference model system. This was achieved by solving the linear Poisson-Boltzmann equation with the MEAD program 17 and using Em(BChlb) = –665 mV and Em(BPheob) = –429 mV (based on Em(BChlb) = –700 mV and Em(BPheob) = –500 mV for one-electron reduction measured in dimethylformamide 18,19), considering the solvation energy difference). The Em(QA) value was calculated, using the reference Em value of –256 mV versus NHE for menaquinone-2 in water 20. The difference in the Em value of the protein relative to the reference system was added to the known Em value. To account for the ensemble of protonation patterns, a Monte Carlo method with Karlsberg was used for sampling 21. The linear Poisson-Boltzmann equation was solved using a three-step grid-focusing procedure with resolutions of 2.5 Å, 1.0 Å, and 0.3 Å. Monte Carlo sampling provided the probabilities [Aox] and [Ared] of the two redox states of molecule A, and Em was evaluated using the Nernst equation. A bias potential was applied to ensure an equal amount of both redox states ([Aox] = [Ared]), thus determining the redox midpoint potential as the resulting bias potential. To ensure consistency with previous computational results, we used identical computational conditions and parameters as previous studies (e.g., 13), performing all computations at 300 K, pH 7.0, and an ionic strength of 100 mM. The dielectric constants were set to 4 for the protein interior and 80 for water.

QM/MM calculations

Earlier version: We employed the restricted DFT method for describing the closed-shell electronic structure and the unrestricted DFT method for the open-shell electronic structure with the B3LYP functional and LACVP* basis sets using the QSite 21 program. To neutralize the entire system, counter ions were added randomly around the protein using the Autoionize plugin in VMD 22. In the QM region, all atom positions were relaxed in the QM region, while the H-atom positions were relaxed in the MM region. The QM regions were defined as follows: for the BChlb pair [PLPM]: the side chains of the ligand residues (His-L173 and His-M200) and H-bod partners (His-L168, Tyr-M195, and Thr-L248); for accessory BChlb: BL/BM and the side chain of the ligand residue (His-L153 for BL/His-M180 for BM); for BPheob: HL/HM.

Revised version: We employed the restricted DFT method for describing the closed-shell electronic structure and the unrestricted DFT method for the open-shell electronic structure with the B3LYP functional and LACVP* basis sets using the QSite 22 program. To neutralize the entire system, counter ions were added randomly around the protein using the Autoionize plugin in VMD 23. In the QM region, all atom positions were relaxed in the QM region, while the H-atom positions were relaxed in the MM region. The QM regions were defined as follows: for the BChlb pair [PLPM]: the side chains of the ligand residues (His-L173 and His-M200) and H-bod partners (His-L168, Tyr-M195, and Thr-L248); for accessory BChlb: BL/BM and the side chain of the ligand residue (His-L153 for BL/His-M180 for BM); for BPheob: HL/HM.

NSD analysis

Earlier version: To analyze the out-of-plane distortions of chlorin rings, we employed an NSD procedure with the minimal basis approximation, where the deformation profile can be represented by the six lowest-frequency normal modes, i.e., ruffling (B1u), saddling (B2u), doming (A2u), waving (Eg(x) and Eg(y)), and propellering (A1u) modes 9, 10. The NSD analysis was performed in the following three steps, as performed previously 6. First, the atomic coordinates of the Mg-substituted macrocycle were extracted from the crystal (or QM/MM optimized) structure. Second, the extracted coordinates were superimposed on the reference coordinates of the macrocycle. The superimposition is based on a least-square method, and the mathematical procedure is described in Ref. 23. Finally, the out-of-plane distortion in the superimposed coordinates was decomposed into the six lowest-frequency normal modes by the projection to the reference normal mode coordinates as where d Γ represents the distortion component of the mode Γ (i.e., Γ = B1u, B2u, A2u, Eg(x), Eg(y), or A1u), Δzi is the z-component of the superimposed coordinates in the ith heavy atom, and is the z-component of the normalized eigenvector of the reference normal mode Γ in the ith heavy atom. N represents the number of heavy atoms. See ref. 6 for further details.

Revised version: To analyze the out-of-plane distortions of chlorin rings, we employed an NSD procedure with the minimal basis approximation, where the deformation profile can be represented by the six lowest-frequency normal modes, i.e., ruffling (B1u), saddling (B2u), doming (A2u), waving (Eg(x) and Eg(y)), and propellering (A1u) modes 9, 10. The NSD analysis was performed in the following three steps, as performed previously 6. First, the atomic coordinates of the Mg-substituted macrocycle were extracted from the crystal (or QM/MM optimized) structure. Second, the extracted coordinates were superimposed on the reference coordinates of the macrocycle. The superimposition is based on a least-square method, and the mathematical procedure is described in Ref. 24. Finally, the out-of-plane distortion in the superimposed coordinates was decomposed into the six lowest-frequency normal modes by the projection to the reference normal mode coordinates as where dΓ represents the distortion component of the mode Γ (i.e., Γ = B1u, B2u, A2u, Eg(x), Eg(y), or A1u), Δzi is the z-component of the superimposed coordinates in the ith heavy atom, and is the z-component of the normalized eigenvector of the reference normal mode Γ in the ith heavy atom. N represents the number of heavy atoms. See ref. 6 for further details.

Results and discussion

Energetically asymmetric electron transfer branches

Earlier version: The XFEL structures show that the Em values for BL are ∼50 mV higher than those for BM, which facilitates the formation of the charge-separated [PLPM] •+BL•– state and thereby electron transfer along the L-branch (Figures 2 and 3). As the Em profile is substantially consistent with the Em profile for PbRC from Rhodobacter sphaeroides 13, it seems plausible that the charge-separated [PLPM] •+BL•– and [PLPM] •+HL•– states in the active L-branch are energetically lower than the [PLPM] •+B •– and [PL PM] •+H •– states in the inactive M-branch, respectively, as demonstrated in QM/MM/PCM calculations 24. Indeed, the calculated Em values are largely correlated with the LUMO levels calculated using a QM/MM approach, as suggested previously (coefficient of determination R2 = 0.98, Figure S1). The Em(HL) value of –597 mV is in line with the experimentally estimated value of ca. –600 mV for HL in PbRC from Blastochloris viridis 25.

Revised version: The XFEL structures show that the Em values for BL are ∼50 mV higher than those for BM, which facilitates the formation of the charge-separated [PLPM] •+BL•- state and thereby electron transfer along the L-branch (Figures 2 and 3). As the Em profile is substantially consistent with the Em profile for PbRC from Rhodobacter sphaeroides 13, it seems plausible that the charge-separated [PLPM] •+BL•– and [PLPM] •+HL•– states in the active L-branch are energetically lower than the [PLPM] •+BM•- and [PLPM] •+HM•– states in the inactive M-branch, respectively, as demonstrated in QM/MM/PCM calculations 25. Indeed, the calculated Em values are largely correlated with the LUMO levels calculated using a QM/MM approach, as suggested previously (coefficient of determination R2 = 0.98, Figure S1). The Em(HL) value of –597 mV is in line with the experimentally estimated value of ca. –600 mV for HL in PbRC from Blastochloris viridis 26.

Em profiles along the L- and M-branches in the XFEL structures for dataset a. (a) 0 ps. (b) 5 ps. (c) 300 ps.

Em profiles along the L- and M-branches in the XFEL structures for dataset b. (a) 0 ps. (b) 1 ps. (c) 5 ps. (d) 20 ps. (e) 300 ps. (f) 8 μs.

Removed in the revised version: Residue pairs that are responsible for Em(BL) > Em(BM).

Earlier version: Among the L/M residue pairs, the Phe-L181/Tyr-M208 pair contributes to Em(BL) > Em(BM) most significantly (27 mV), facilitating L branch electron transfer, as suggested in theoretical studies 26 (Table 1, Figure 3a). This result is also consistent with the contribution of the Phe-L181/Tyr-M210 pair to the difference between Em(BL) and Em(BM), which was the largest in PbRC from Rhodobacter sphaeroides 27 (26 mV 13). The Asn-L158/Thr-M185 pair also contributes to the difference between Em(BL) and Em(BM) (11 mV, Table 1), as does the Val-L157/Thr-M186 pair in PbRC from Rhodobacter sphaeroides (22 mV 13).

Revised version: Among the L/M residue pairs, the Phe-L181/Tyr-M208 pair contributes to Em(BL) > Em(BM) most significantly (27 mV), facilitating L branch electron transfer, as suggested in theoretical studies 27 (Table 1, Figure 3a). This result is also consistent with the contribution of the Phe-L181/Tyr-M210 pair to the difference between Em(BL) and Em(BM), which was the largest in PbRC from Rhodobacter sphaeroides 28 (26 mV 13). The Asn-L158/Thr-M185 pair also contributes to the difference between Em(BL) and Em(BM) (11 mV, Table 1), as does the Val-L157/Thr-M186 pair in PbRC from Rhodobacter sphaeroides (22 mV 13).

Earlier version: The Em values for HL are ∼50 mV higher than those for HM in the dark state and [5 ps and 300 ps] XFEL structures, as observed in Em(BL) and Em(BM) (Figure 2a,c,e). However, the Em difference decreases to ∼25 mV in the [1 ps, 20 ps, and 8 μs] XFEL structures (Figure 2b,d,f), which implies that the dark state and [5 ps and 300 ps] XFEL structures are distinct from the [1ps, 20 ps, and 8 μs] XFEL structures (see below). Below, we discuss the dark state structure if not otherwise specified.

Revised version: The Em values for HL are ∼50 mV higher than those for HM in the dark state and [5 ps and 300 ps] XFEL structures, as observed in Em(BL) and Em(BM) (Figure 2a,c,e). However, the Em difference decreases to ∼25 mV in the [1 ps, 20 ps, and 8 μs] XFEL structures (Figure 2b,d,f), which implies that the dark state and [5 ps and 300 ps] XFEL structures are distinct from the [1 ps, 20 ps, and 8 μs] XFEL structures (see below). Below, we discuss the dark state structure if not otherwise specified.

The Ala-L120/Asn-M147 pair contributes to Em(HL) > Em(HM) most significantly (38 mV) (Table 2, Figure S2). However, this holds true only for PbRC from Blastochloris viridis, as Asn-M147 is replaced with alanine (Ala-M149) in PbRC from Rhodobacter sphaeroides. The Asp-L218/Trp-M252 pair decreases Em(HM) with respect to Em(HL), thereby contributing to Em(HL) > Em(HM) (20 mV) (Table 2, Figure S2). Arg-L103 orients toward the protein interior, whereas Arg-M130 orients toward the protein exterior (Figure S2), which contributes to Em(HL) > Em(HM) (17 mV) (Table 2). Ser-M271 forms an H-bond with Asn-M147 near HM (Figure 3b). Thus, the contribution of Ser-M271 to Em(HL) is large, although this residue is replaced with alanine (Ala-M273) in PbRC from Rhodobacter sphaeroides.

Relevance of structural changes observed in XFEL structures

Earlier version: According to Dods et al., the 5-ps and 20-ps structures correspond to the charge-separated [PLPM]•+HL•– state 8. If this is the case, Em(HL) is expected to be exclusively higher in the 5-ps and 20-ps structures than in the other XFEL structures due to the stabilization of the [PLPM]•+H •– state by protein reorganization. In dataset a, the E (H) value is only 4 mV higher in the 5-ps structure than in the dark structure (Figure 5a). In dataset b, the Em(HL) value is ∼20 mV higher in the 5- and 20-ps structures than in the dark structure (Figure 5b). However, the Em(HL) value is 25 mV higher in the 300-ps structure than in the dark structure. Tables 3 and 4 show the residues that contribute to the slight increase in Em(HL) most significantly in the 5- and 20-ps structures. Most of these residues were in the region where Dods et al. specifically performed multiple rounds of partial occupancy refinement (e.g., 153–178, 190, 230 and 236–248 of subunit L and 193–221, 232, 243–253, 257–266 of subunit M) 8. In dataset b (Table 4), which has more data points than dataset a (Table 3), the contributions of these residues to Em(HL) often fluctuate (e.g., upshift/downshift followed by downshift/upshift) at different time intervals (e.g., 1 to 5 ps, 5 to 20 ps, and 20 to 300 ps). This result suggests that the structural differences among the XFEL structures are not related to the actual time course of charge separation. Furthermore, the Em(HM) value in the inactive M branch is also ∼15 mV higher in the 5- and 20-ps structures than in the dark structure (Figure 5b). These results suggest that the ∼20 mV higher Em(HL) value in the 5- and 20-ps structures is not specifically due to the formation of the [PLPM]•+HL•– state. Thus, the stabilization of the [PLPM]•+HL•– state owing to protein reorganization is not clearly observed in the Em(HL) values.

Revised version: According to Dods et al., the 5-ps and 20-ps structures correspond to the charge-separated [PLPM]•+HL•– state 8. If this is the case, Em(HL) is expected to be exclusively higher in the 5-ps and 20-ps structures than in the other XFEL structures due to the stabilization of the [PLPM]•+HL•– state by protein reorganization. In dataset a, the Em(HL) value is only 4 mV higher in the 5-ps structure than in the dark structure (Figure 5a). In dataset b, the Em(HL) value is ∼20 mV higher in the 5- and 20-ps structures than in the dark structure (Figure 5b). However, the Em(HL) value is 25 mV higher in the 300-ps structure than in the dark structure. Tables 3 and 4 show the residues that contribute to the slight increase in Em(HL) most significantly in the 5- and 20-ps structures. Most of these residues were in the region where Dods et al. specifically performed multiple rounds of partial occupancy refinement (e.g., 153–178, 190, 230 and 236–248 of subunit L and 193–221, 232, 243– 253, 257–266 of subunit M) 8. In dataset b (Table 4), which has more data points than dataset a (Table 3), the contributions of these residues to Em(HL) often fluctuate (e.g., upshift/downshift followed by downshift/upshift) at different time intervals (e.g., 1 to 5 ps, 5 to 20 ps, and 20 to 300 ps). This result suggests that the structural differences among the XFEL structures are not related to the actual time course of charge separation. Furthermore, the Em(HM) value in the inactive M branch is also ∼15 mV higher in the 5- and 20-ps structures than in the dark structure (Figure 5b). These results suggest that the ∼20 mV higher Em(HL) value in the 5- and 20-ps structures is not specifically due to the formation of the [PLPM]•+HL•– state. Thus, the stabilization of the [PLPM]•+HL•– state owing to protein reorganization is not clearly observed in the Em(HL) values.

Added in the revised version: Residue pairs that are responsible for Em(BL) > Em(BM).

Time-dependent Em changes for BChlb and BPheob in the XFEL structures. (a) Dataset a. (b) Dataset b. ΔEm denotes the Em shift with respect to the dark state structure. Black solid lines: PL; black dotted lines: PM; blue solid lines BL; blue dotted lines: BM; red solid lines: HL; red dotted lines: HM.

A normal-coordinate structural decomposition (NSD) analysis 9, 10 of the out-of-plane distortion of the chlorin ring is sensitive to subtle structural changes in the chlorin ring, which are not distinct in the Em changes 6. QM/MM calculations indicate that HL•- formation induces the saddling mode in the chlorin ring, which describes the movement of rings I and III being in the opposite direction to the movement of rings II and IV along the normal axis of the chlorin ring (Tables 5 and 6). However, (i) in the XFEL structures, the saddling mode of HL remains practically unchanged in dataset a during electron transfer (Figure 6 and Tables S2 and S3). In dataset b, the saddling mode of HL is induced most significantly at 1 ps, which does not correspond to the charge-separated [PLPM]•+HL•– state (Figure 7). (ii) In addition, the ruffling mode is more pronounced than the saddling mode in HL (Figure 7), which suggests that the observed deformation of HL is not directly associated with the reduction of HL.

Time-dependent changes in the lowest frequency out-of-plane modes of the chlorin rings in the XFEL structures (dataset a). Sad: saddling (red); ruf: ruffling (blue); dom: doming (green); wav(x, y): waving (x, y) (gray, dark blue); pro: propellering (orange). Solid and dotted lines indicate L and M branches, respectively. See Table S2 for the absolute values in the dark state for dataset a.

Time-dependent changes in the lowest frequency out-of-plane modes of the chlorin rings in the XFEL structures (dataset b). Sad: saddling (red); ruf: ruffling (blue); dom: doming (green); wav(x, y): waving (x, y) (gray, dark blue); pro: propellering (orange). Solid and dotted lines indicate L and M branches, respectively. See Table S3 for the absolute values in the dark state for dataset b.

One might argue that the loss of the link between the formation of the charge-separated state and the Em(HL) change (Figure 5) is not due to experimental errors but rather represents the actual ps-timescale phenomena during the primary charge-separation reactions (e.g., Dods et al. noted that “the primary electron-transfer step to HL is more rapid than conventional Marcus theory” 8). However, even if this were the case, this hypothesis regarding the relevance of the XFEL structures to the electron-transfer events could be further explored by examining the changes in Em(QA) among the XFEL structures, considering the relatively slow electron-transfer step to QA that allows sufficient protein relaxation to occur (e.g., Dods et al. stated that “the electron-transfer step to QA has a single exponential decay time of 230 ± 30 ps, consistent with conventional Marcus theory” 8). That is, if the Em(QA) values are not higher in the 300-ps and 8-μs structures than in the other structures, it suggests that significant experimental errors exist, rendering the XFEL structures irrelevant to the electron transfer events. Consistent with this perspective, the present results demonstrate that the Em(QA) values in the 300-ps and 8-μs structures are not significantly higher than those in the other structures, including the dark state structure (Figure 8). Consequently, the lack of a clear relationship between the charge separated state and the changes in Em(QA) at 300 ps and 8-μs further strengthens the argument that the XFEL structures are irrelevant to the electron transfer events.

Earlier version: Time-dependent Em changes for QA in the XFEL structures. (a) Dataset a. (b) Dataset b. ΔEm denotes the Em shift with respect to the dark state structure.

Revised version: Time-dependent Em changes for QA in the XFEL structures. (a) Dataset a. (b) Dataset b. ΔEm denotes the Em shift with respect to the dark state structure. Note that the calculated Em(QA) values for dataset a and dataset b in the dark structure are –223 mV and –209 mV, respectively, which are comparable to experimentally measured values of –150 mV for PbRC from Blastochloris viridis (menaquinone) 29 and –180 mV for PbRC from Rhodobacter sphaeroides (ubiquinone) 30.

Earlier version: In summary, the Em values in the active L branch are higher than those in the inactive M branch in the XFEL structures, which suggests that electron transfer via BL•– and HL•– is energetically more favored than that via BM•– and HM•– (Figure 2). The Phe-L181/Tyr-M208 pair contributes to the difference between Em(BL) and Em(BM) the most significantly, as observed in the Phe-L181/Tyr-M210 pair in PbRC from Rhodobacter sphaeroides 13, 27. The stabilization of the [PLPM]•+HL•– state owing to protein reorganization is not clearly observed in the Em(HL) values (Figure 5). The absence of the induced saddling mode in the HL chlorin ring in the 5- and 20-ps structures suggests that HL•– does not specifically exist in these XFEL structures (Figures 6 and 7). The cyclic fluctuations in the contributions of the residues to Em(HL) at different time intervals suggest that the structural differences among the XFEL structures are not related to the actual time course of charge separation (Table 4). The major limitation of the structural studies conducted by Dods et al. 8 is the relatively low resolution of their XFEL structures, primarily at 2.8 Å. Consequently, the observed changes in Em values and chlorin ring deformations are more likely to reflect experimental errors rather than actual structural changes induced by electron transfer events. This concern is reinforced by the lack of a clear relationship between the actual QA•– formation and the Em(QA) values in the 300-ps and 8-μs structures (Figure 8). Consequently, the proposed time-dependent structural changes proposed by Dods et al. 8 are highly likely irrelevant to the electron transfer events.

Revised version: In summary, the Em values in the active L branch are higher than those in the inactive M branch in the XFEL structures, which suggests that electron transfer via BL•- and HL•- is energetically more favored than that via BM•- and HM•– (Figure 2). The Phe-L181/Tyr-M208 pair contributes to the difference between Em(BL) and Em(BM) the most significantly, as observed in the Phe-L181/Tyr-M210 pair in PbRC from Rhodobacter sphaeroides 13, 28. The stabilization of the [PLPM]•+HL•– state owing to protein reorganization is not clearly observed in the Em(HL) values (Figure 5). The absence of the induced saddling mode in the HL chlorin ring in the 5- and 20-ps structures suggests that HL•– does not specifically exist in these XFEL structures (Figures 6 and 7). The cyclic fluctuations in the contributions of the residues to Em(HL) at different time intervals suggest that the structural differences among the XFEL structures are not related to the actual time course of charge separation (Table 4). The major limitation of the structural studies conducted by Dods et al. 8 is the relatively low resolution of their XFEL structures, primarily at 2.8 Å. Consequently, the observed changes in Em values and chlorin ring deformations are more likely to reflect experimental errors rather than actual structural changes induced by electron transfer events. This concern is reinforced by the lack of a clear relationship between the actual QA•- formation and the Em(Qm) values in the 300-ps and 8-μs structures (Figure 8). Consequently, the proposed time-dependent structural changes proposed by Dods et al. 8 are highly likely irrelevant to the electron transfer events.

Hence, it is crucial to exercise caution when interpreting time-dependent XFEL structures, especially in the absence of comprehensive evaluations of the energetics and accompanying structural changes. This cautionary note should serve as a counterargument in the future, highlighting the potential pitfalls associated with presenting time-dependent XFEL structures of insufficient quality and drawing conclusive interpretations of protein structural changes that may not be distinguishable from significant experimental errors. Future high-resolution structures may provide further insights into the actual structural changes relevant to electron transfer events. By combining both high-resolution structures and rigorous energetic evaluations, a more comprehensive understanding of the protein structure-function relationship can be achieved.

Supporting information